<br><h3><b>CITS2402 Lab 1</b></h3><br>

# Diversity in Study: What (Else) Are We All Studying?

----


## Data Science Across the Disciplines

Data Science has application right across the disciplines. In this course we'll see case studies ranging from the Environment to Economics to Health.

Before we start this journey, it's interesting to see what diverse areas are studied by data science students. As we said in the Welcome lecture, there are many students in the unit with many different study backgrounds and paths.

To get a feel for this, we'll look at the range of enrolment patterns.

Note that some randomisation has been applied to the enrolment patterns to ensure privacy and anonymity. However the range of disciplines is indicative of the unit.

In this lab we will also:

- practise working with text files, strings and lists
- practise writing user-defined functions
- do our first data cleaning
- do some data manipulation to observe features of our data
- explain how checked lab questions will be assigned and marked

#### Reference links

We'll be using standard python - no additional modules are required (or should be used). The latest documentation will be useful, and it is suggested you bookmark it for future use.

* _The Python Language Reference_, https://docs.python.org/3/reference/
* _The Python Standard Library_, https://docs.python.org/3/library/
* _The Python Tutorial_, https://docs.python.org/3/tutorial/

## Data Acquisition

In future labs we'll be hunting down data as part of the work. In this case however the data will be provided in a text file. Typically the next task after acquiring the data is cleaning it and putting it into a more useable format, usually requiring some string processing. We'll have some practice at this in this lab.

* Open the directory containing this lab sheet from the folder 📂 icon. Click on the file `unit-patterns.txt` to inspect the file.

The file contains all the different study plans of a cohort of Data Science students with some obfuscation to ensure privacy and anonymity.

This file was output from another system, and is in the raw format from that system. This is a *tab-separated values (tsv)* file with two *fields*.

The first field contains a set of units studied (separated by plus signs), the second (after the tab character) contains a number and a percentage.

Note that in this lab we will refer to an instance of a unit code combined with a teaching period (eg a semester) as a '_unit'_ - so for example, CITS1401-1 and CITS1401-2 would be treated as two distinct units. These units have the same alphanumeric _unit code_ (CITS1401).

---

Let's start with a *constant* to identify the file. By convention we'll use UPPERCASE for constants.

Run the following cell to make the assignment.

```python
DATAFILE = "unit-patterns.txt"
```

As we saw in the _Getting Started_ lab, we can check to see that the assignment was made just by typing the variable name as the last thing in the cell, and then running the cell. (This is because it prints out the results of whatever was the last evaluation, and a variable simply evaluates to its contents.)

- Try this now.

```python
print(DATAFILE)

```

Output:
```
unit-patterns.txt
```

#### Reading the file

* Use the built-in function `open()` and the text file IO method `readline()` to read in and print out the first 5 lines of the file.
* Use python's `with` construct so that you don't need to worry about closing the file.

Your output should start like this:
```
 CITS2224-1 + CITS2023-2 + CITS3221-2 + CITS3224-1 + CITS3225-1 	 1 (2.0%) 

 CITS2023-2 + ENSC3221-2 + ENSC3224-1 + ENSC3225-1 + ENSC3227-2 + ENSC3210-1 + ENSC3221-1 + MECH3023-2 	 1 (2.0%) 

 CITS1021-2 + CITS2021-1 + CITS2023-2 + ENSC2224-1 + ENSC3211-2 + ENSC3213-2 + MATH2221-1 + STAT1524-1 	 1 (2.0%) 
...
```

Tip: If you're a little rusty on reading files, or haven't used `with` before, you will find this tutorial useful: [7.2. Reading and Writing Files](https://docs.python.org/3/tutorial/inputoutput.html?highlight=readline#reading-and-writing-files).

Also the Case Study "FakeNews" provides a couple of examples of working with a `.tsv` file.



```python
with open(DATAFILE,'r') as file:
     count=1
     for i in file:
         if count<6:
            print(i)
            count+=1



```

Output:
```
CITS2224-1 + CITS2023-2 + CITS3221-2 + CITS3224-1 + CITS3225-1 	 1 (2.0%) 

 CITS2023-2 + ENSC3221-2 + ENSC3224-1 + ENSC3225-1 + ENSC3227-2 + ENSC3210-1 + ENSC3221-1 + MECH3023-2 	 1 (2.0%) 

 CITS1021-2 + CITS2021-1 + CITS2023-2 + ENSC2224-1 + ENSC3211-2 + ENSC3213-2 + MATH2221-1 + STAT1524-1 	 1 (2.0%) 

 CITS2023-2 + CITS3220-2 + CITS3223-2 + CITS3021-1 + CITS3025-1 + MATH3224-1 + MATH3221-1 + STAT3264-1 + STAT3260-2 	 1 (2.0%) 

 BIOC3221-1 + BIOC3224-1 + BIOC3225-2 + BIOC3225-2 + CITS1023-2 + CITS2021-1 + CITS2023-2 + MATH1211-1 	 1 (2.0%)
```

## Marked Questions and Answers

Many of the labs this semester will have one or two checked (or 'autograded') answers. These will contribute to the practical mark for the semester.

The checked answers will generally have three parts (or three cells):

#### 1. A specification (the question)

This will ask you to write some code, normally in the form of a function. The question will specify (in the context of what we've done in the build up to that question in the lab sheet) what the function should do, usually including its inputs and outputs.

#### 2. An answer cell

The first code cell is where you enter your code to be checked. This may contain a "stub" - that is, a part of the code to be completed. Often it will be a function header. For example:
```
def get_patterns(filename):
```
Your task is to complete the code, according to the specification.


#### 3. Test cells

The second code cell is where the code to check the answer will go and is not editable. It may be blank, but will usually contain some example tests that you can use as sample tests for your code. _There will be additional tests applied that are not included in the sample tests._ They may use slightly different data. In some cases manual checks will also be made.

You can run the sample tests as many times as you like. It should be noted that passing the sample tests does NOT mean that your code is correct or will pass the additional tests. The goal is to meet the specification, not the sample tests, which are provided to help you get on the right track.

You should write your own *unit tests* to test your code extensively, like you would be required to do in a work situation. You should consider, for example, what would happen with a different set of data (of the same format).

### Collection and deadlines

The notebook will be collected (from the same directory it was distributed to) and marked, by running the tests, after the lab deadline. (As the same notebook will be collected, you may wish to take a backup copy of the notebook before starting work. Alternatively you can use the TimeTravel feature to revert the notebook if you need to start again.)

It is intended that the lab is done in the week it is assigned. However, everyone will automatically receive a one week "extension" from the week of the lab to finish off the work if necessary and to account for any unforeseen circumstances and special considerations. This is a 'timed assessment' and after the 7 days there are no further extensions. This is needed so that we can discuss and answer any questions arising from the lab.

The last deadline for credit will therefore be **11.59pm on the Friday of the week after the lab is released**.

---

Lets try a practice exmaple.



#### Checked Solution Practice [1 point] 

#### Problem Specification

Complete the following stub to create a function `get_count(unitcode, filename)` that:

- takes as input a unit code, such as CITS1021, and a filename for a text file in tsv format as above
- returns as output the number of students who are enrolled in that unit code (in either semester)

_Hint: In case you're a bit rusty, we cover a  similar example in the Fake News lecture._

---

#### Answer cell

Complete your code here. You should remove the `raise NotImplementedError()` when you put in your own code.

def get_count (unitcode,filename):
    nospace_list=[]
    with open(filename,"r") as file:
        data=file.read().splitlines()
        print(data)

        counts=0
        for row in data:
            if unitcode in row:
                counts+=1
        print(counts)

get_count('STAT2021',DATAFILE)

<!-- grade_id: 0c596f, points: 0, grade: false, solution: true -->
```python
def get_count (unitcode,filename):
    nospace_data=nospace(filename)
    counts=0
    
    for row in nospace_data:
        for unit in row:
            if unitcode == unit[:-2]:  #remove the last 2 character starting from dash
                counts +=1
    return counts




#print(get_count('STAT2021',DATAFILE))

```

Try out your function. Add some print statements in a new cell and see if you get what you would expect.

_Note: You should avoid writing to files or printing large amounts within your submitted checked function (if not asked to) as this may cause problems for the autograder. Use a separate cell for your testing._

```python
x = "apple-2"
x[:-2]
```

Output:
```
'apple'
```

#### Running the tests

The test cells will often include some sample tests that you can run to see if your code passes. There will also be some additional tests that are only run once the code is collected after the deadline.

If your code runs successfully it will simply finish as usual (without throwing an exception).

If your code does not pass the sample tests, you will normally get an error message. This might be something relatively simple (you can ignore the middle bit) like:

```
---------------------------------------------------------------------------
AssertionError                            Traceback (most recent call last)
/tmp/ipykernel_26920/1618265836.py in <module>
      1 from nose.tools import assert_equal
      2 DATA = "unit-patterns.txt"
----> 3 assert_equal(get_count('STAT2021',DATAFILE), 72)
      4 assert_equal(get_count('STAT2023',DATAFILE), 67)
      5 print("If you didn't get an error, then so far so good - your code passed the sample test.")
/usr/lib/python3.8/unittest/case.py in assertEqual(self, first, second, msg)
    910         """
    911         assertion_func = self._getAssertEqualityFunc(first, second)
--> 912         assertion_func(first, second, msg=msg)
    913 
    914     def assertNotEqual(self, first, second, msg=None):
/usr/lib/python3.8/unittest/case.py in _baseAssertEqual(self, first, second, msg)
    903             standardMsg = '%s != %s' % _common_shorten_repr(first, second)
    904             msg = self._formatMessage(msg, standardMsg)
--> 905             raise self.failureException(msg)
    906 
    907     def assertEqual(self, first, second, msg=None):
AssertionError: 65 != 72
```

This shows which test failed (the first arrow) and, at the bottom, what the incorrect answer was. Here it was expecting 72 but the function returned 65.

It might be something more complex looking, but may still give you some clues. For example, the problem here is extra whitespace after the entry 'CITS2224-1':

```
---------------------------------------------------------------------------
AssertionError                            Traceback (most recent call last)
<ipython-input-82-f1838aea55b1> in <module>
      1 from nose.tools import assert_equal
      2 DATA = "unit-patterns.txt"
----> 3 assert_equal(get_patterns(DATA)[0],['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1'])
      4 assert_equal(get_patterns(DATA)[5],['CITS2021-1', 'CITS2023-2', 'GRMN2025-1', 'GRMN2020-2', 'MATH1214-1', 'MATH2521-2', 'PHYS2221-1', 'PHYS2223-2'])
...
668     def fail(self, msg=None):
    669         """Fail immediately, with the given message."""
--> 670         raise self.failureException(msg)
    671 
    672     def assertFalse(self, expr, msg=None):
AssertionError: Lists differ: ['CITS2224-1 ', 'CITS2023-2 ', 'CITS3221-2 ', 'CITS3224-1 ', 'CITS3225-1 '] != ['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1']

First differing element 0:
'CITS2224-1 '
'CITS2224-1'

- ['CITS2224-1 ', 'CITS2023-2 ', 'CITS3221-2 ', 'CITS3224-1 ', 'CITS3225-1 ']
?             -              -              -              -              -

+ ['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1']
```

The error messages may vary in their usefulness - they are not intended as a sole debugging tool, but often they will give some pointer as to what went wrong. In this case it is fairly clear on closer inspection that there is whitespace that shouldn't be there.

You will likely need to come up with your own unit tests to figure out what is wrong. You can then correct your code and try again.



## Unit Testing

In general, you should create your own tests for your code before submitting it.

[TechTarget](https://www.techtarget.com/searchsoftwarequality/definition/unit-testing) provides a good definition of unit testing from our perspective:

> Unit testing is a software development process in which the smallest testable parts of an application, called units, are individually and independently scrutinized for proper operation. This testing methodology is done during the development process by the software developers and sometimes QA staff.  The main objective of unit testing is to isolate written code to test and determine if it works as intended.
>
> Unit testing is an important step in the development process, because if done correctly, it can help detect early flaws in code which may be more difficult to find in later testing stages.

In this course we will use a compositional approach (or, looking from the other direction, _problem decomponsition_) for problem solving (and programming). This means we will use functions (or "functional units") which each have a specific, well-defined, purpose, and compose those to form a problem solution. As the above quote suggests, it is much easier to find bugs by testing a stand-alone function than it is once it is composed into a bigger program.

**Examples**

So what might a unit test look like? Normally it would be something where you can verify that the code does what you expect. In this case, one example would be to simply pick a course where you can easily count the occurrences yourself by searching through the input file. Something like CITS1021 occurs many times, so its easy to get the (manual) count wrong. So try something that is a bit rarer.

* Check your code on MUSC2275. Use 'find' in the text file to manually check that the number returned is correct.

**Edge Cases (or Limit Cases)**

Another common way of checking code is to check for "_edge cases_" - cases at the extremes. For this example, one edge case would be a unit code that doesn't appear at all. You should check that this returns zero. Another would be to write code to find the unit code that appears in the largest number of student study plans, and check that this is not greater than the total number of students (that is, lines in the input). If it is, you know there must be something wrong with the code.

* Write unit tests for each of the above test cases.

**Reductio ad Absurdum?**

When testing (and debugging) code, a useful technique is to ask what is the smallest example you can test it (or break it - see below) using. Let's say we want to sort the list of study patterns. Its very difficult for a human to check the answer when there may be hundreds of units or unit sets. However entirely equivalent (from a sorting perspective, not a cleaning perspective) cases can be designed just using letters such as `a`, `b`, `c`,... 

**White Hat mentality**

Another way of thinking about unit tests is to think about intentionally trying to "break" your code.  What cases might cause it difficulty? In the sorting example, can you come up with combinations of the letters that will in which your sort will fail?

In fact, on a larger scale, this is what a "White Hat Hacker" does, and can make you [a lot of money](https://totalcoin.io/articles/183-24-500-how-do-white-hat-hackers-cash-in-on-crypto.html)!

---

The following is an example of a Test Cell for the question above, with a couple of sample tests shown. Try running the code.


#### Test cell


<!-- grade_id: get_count, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal  # ignore this line
DATA = "unit-patterns.txt"

# Sample tests
assert_equal(get_count('STAT2021',DATA), 72)
assert_equal(get_count('STAT2023',DATA), 67)
print("If you didn't get an error, then so far so good - your code passed the sample tests.")

```

Output:
```
/tmp/ipykernel_3914/3304731750.py in get_count(unitcode, filename)
      1 def get_count (unitcode,filename):
----> 2     nospace_data=nospace(filename)
      3     counts=0
      4 
      5     for row in nospace_data:

/tmp/ipykernel_3914/1014921456.py in nospace(filename)
      4 
      5     with open(filename,'r') as file:
----> 6         for content in pattern_lists:
      7             newlist=[] #refresh the list to start over
      8             for units in content:

NameError: name 'pattern_lists' is not defined
```

If your code passed the above test, try it again with some deliberate errors in your code to see what happens. Then make sure you put it back to the correct version.

When the lab is actually collected, additional tests will be applied. In this case we'll reveal that there will be two additional tests - one for MUSC2275, and one corresponding to one of the edge-case unit tests described above.

(You can see why the marked questions must have additional "hidden" tests. Otherwise we could just look at the tests above, and write the following code which would also pass!)

```
def get_count (unitcode, filename):
    if unitcode == 'STAT2021':
        return 72
    else:
        return 67
```

---


Now lets return to our case study...

#### From Files to Lists

* Write some code to read the data file again, but this time, as you read the lines from the file, store them in a list called `pattern_strings`. If you're a bit rusty on adding things to a list, you might wish to review Chapter 10 of *Think Python 2e*.

Printing the first 5 items of your list should give:
```
[' CITS2224-1 + CITS2023-2 + CITS3221-2 + CITS3224-1 + CITS3225-1 \t 1 (2.0%) \n', ' CITS2023-2 + ENSC3221-2 + ENSC3224-1 + ENSC3225-1 + ENSC3227-2 + ENSC3210-1 + ENSC3221-1 + MECH3023-2 \t 1 (2.0%) \n', ' CITS1021-2 + CITS2021-1 + CITS2023-2 + ENSC2224-1 + ENSC3211-2 + ENSC3213-2 + MATH2221-1 + STAT1524-1 \t 1 (2.0%) \n', ' CITS2023-2 + CITS3220-2 + CITS3223-2 + CITS3021-1 + CITS3025-1 + MATH3224-1 + MATH3221-1 + STAT3264-1 + STAT3260-2 \t 1 (2.0%) \n', ' BIOC3221-1 + BIOC3224-1 + BIOC3225-2 + BIOC3225-2 + CITS1023-2 + CITS2021-1 + CITS2023-2 + MATH1211-1 \t 1 (2.0%) \n']
```

Rather than using `readline()` or other read methods, use an *iterator* on the file directly to get each line in turn. (Notice that the iterator retains the newline character.)

```python
pattern_strings =[]
with open(DATAFILE,'r') as file:
    pattern_strings=file.readlines()
    print(pattern_strings)

```

## Data Cleaning

In this case study we are only interested in the patterns in the first field (to the left of the tab) in the data file. Since the lines contain two data fields separated by a delimiter, the string `partition()` method (or function) is ideal for the job.

* Using the string `partition()` method, modify your code so that `pattern_strings` contains only the unit patterns. You should _partition around the delimiter_ (tab or `'\t'`), not the contents of the fields.

The documentation for [String Methods](https://docs.python.org/3/library/stdtypes.html#string-methods) may be helpful.

Printing the first 5 lines of your list (`print(pattern_strings[:5])`) should now give:

```
[' CITS2224-1 + CITS2023-2 + CITS3221-2 + CITS3224-1 + CITS3225-1 ', ' CITS2023-2 + ENSC3221-2 + ENSC3224-1 + ENSC3225-1 + ENSC3227-2 + ENSC3210-1 + ENSC3221-1 + MECH3023-2 ', ' CITS1021-2 + CITS2021-1 + CITS2023-2 + ENSC2224-1 + ENSC3211-2 + ENSC3213-2 + MATH2221-1 + STAT1524-1 ', ' CITS2023-2 + CITS3220-2 + CITS3223-2 + CITS3021-1 + CITS3025-1 + MATH3224-1 + MATH3221-1 + STAT3264-1 + STAT3260-2 ', ' BIOC3221-1 + BIOC3224-1 + BIOC3225-2 + BIOC3225-2 + CITS1023-2 + CITS2021-1 + CITS2023-2 + MATH1211-1 ']
```

* How many students were enrolled at the time the list was taken? (Single line of code.)



```python
pattern_strings =[]
with open(DATAFILE,'r') as file:
     data=file.readlines()
     #print(data)
     for row in data:
         unit_data=row.partition('\t')
         unit_only=unit_data[0]
         pattern_strings.append(unit_only)

     print(pattern_strings)

```

#### Strings to lists

Rather than the units in each line being part of a string, it is more useful to have them in a list. This will enable us to iterate through the list later. Since there are varying numbers of items, the `split()` method makes sense here, with plus ('`+`') as the delimiter.

* Use the `split` method to split each line into a list of units. Store the result in a variable `pattern_lists`.

This should result in a _list of lists_. Your output should start like this:

```
[[' CITS2224-1 ', ' CITS2023-2 ', ' CITS3221-2 ', ' CITS3224-1 ', ' CITS3225-1 '], [' CITS2023-2 ', ' ENSC3221-2 ', ' ENSC3224-1 ', ' ENSC3225-1 ', ' ENSC3227-2 ', ' ENSC3210-1 ', ' ENSC3221-1 ', ' MECH3023-2 '], [' CITS1021-2 ', ' CITS2021-1 ', ' CITS2023-2 ', ' ENSC2224-1 ', ' ENSC3211-2 ', ' ENSC3213-2 ', ' MATH2221-1 ', ' STAT1524-1 '], [' CITS2023-2 ', ' CITS3220-2 ', ' CITS3223-2 ', ' CITS3021-1 ', ' CITS3025-1 ', ' MATH3224-1 ', ' MATH3221-1 ', ' STAT3264-1 ', ' STAT3260-2 '], [' BIOC3221-1 ', ' BIOC3224-1 ', ' BIOC3225-2 ', ' BIOC3225-2 ', ' CITS1023-2 ', ' CITS2021-1 ', ' CITS2023-2 ', ' MATH1211-1 ']]
```



```python
pattern_lists=[]
pattern_strings =[]
with open(DATAFILE,'r') as file:
     data=file.readlines()
     #print(data)
     for row in data:
         unit_data=row.partition('\t')
         unit_only=unit_data[0].split('+')
         pattern_lists.append(unit_only)

     print(pattern_lists)

```

Output:
```
[[' CITS2224-1 ', ' CITS2023-2 ', ' CITS3221-2 ', ' CITS3224-1 ', ' CITS3225-1 '], [' CITS2023-2 ', ' ENSC3221-2 ', ' ENSC3224-1 ', ' ENSC3225-1 ', ' ENSC3227-2 ', ' ENSC3210-1 ', ' ENSC3221-1 ', ' MECH3023-2 '], [' CITS1021-2 ', ' CITS2021-1 ', ' CITS2023-2 ', ' ENSC2224-1 ', ' ENSC3211-2 ', ' ENSC3213-2 ', ' MATH2221-1 ', ' STAT1524-1 '], [' CITS2023-2 ', ' CITS3220-2 ', ' CITS3223-2 ', ' CITS3021-1 ', ' CITS3025-1 ', ' MATH3224-1 ', ' MATH3221-1 ', ' STAT3264-1 ', ' STAT3260-2 '], [' BIOC3221-1 ', ' BIOC3224-1 ', ' BIOC3225-2 ', ' BIOC3225-2 ', ' CITS1023-2 ', ' CITS2021-1 ', ' CITS2023-2 ', ' MATH1211-1 '], [' CITS2021-1 ', ' CITS2023-2 ', ' GRMN2025-1 ', ' GRMN2020-2 ', ' MATH1214-1 ', ' MATH2521-2 ', ' PHYS2221-1 ', ' PHYS2223-2 '], [' CITS1221-1 ', ' CITS1021-1 ', ' CITS1023-2 ', ' CITS2023-2 ', ' MGMT1135-1 ', ' MGMT1136-2 ', ' STAT1024-1 ', ' STAT2023-2 '], [' CITS1221-2 ', ' CITS2023-2 ', ' CITS3223-2 ', ' CITS3021-1 ', ' CITS3025-1 ', ' CLAN3224-1 ', ' STAT2021-1 ', ' STAT32
```

#### Stripping whitespace

Finally, notice that we've been left with unnecessary whitespace around each unit name.

* Find a suitable string method from the documentation, and modify your code to strip off the extra whitespace.

You should find the most appropriate method. (Don't, for example, use `replace()`.) It is also good practice to use a method that is as "_robust_" as possible to errors in data - for example, that will still work if there is more than one space. (Don't, for example, use indexing and rely on the position of the space.)

Your first 5 patterns should now look like this:
    
```
[['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1'], ['CITS2023-2', 'ENSC3221-2', 'ENSC3224-1', 'ENSC3225-1', 'ENSC3227-2', 'ENSC3210-1', 'ENSC3221-1', 'MECH3023-2'], ['CITS1021-2', 'CITS2021-1', 'CITS2023-2', 'ENSC2224-1', 'ENSC3211-2', 'ENSC3213-2', 'MATH2221-1', 'STAT1524-1'], ['CITS2023-2', 'CITS3220-2', 'CITS3223-2', 'CITS3021-1', 'CITS3025-1', 'MATH3224-1', 'MATH3221-1', 'STAT3264-1', 'STAT3260-2'], ['BIOC3221-1', 'BIOC3224-1', 'BIOC3225-2', 'BIOC3225-2', 'CITS1023-2', 'CITS2021-1', 'CITS2023-2', 'MATH1211-1']]
```

```python
[[1,2,3,4,5,6], [7,8,9,10,11,12]]
#refresh the list as it goes before appending
```

```python
def nospace(filename):

    nospace_list=[]

    with open(filename,'r') as file:
        for content in pattern_lists:
            newlist=[] #refresh the list to start over 
            for units in content:
                newcontent=units.strip()
                newlist.append(newcontent)
            nospace_list.append(newlist)
        return nospace_list
    
    
nospace(DATAFILE)
```

Output:
```
[['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1'],
 ['CITS2023-2',
  'ENSC3221-2',
  'ENSC3224-1',
  'ENSC3225-1',
  'ENSC3227-2',
  'ENSC3210-1',
  'ENSC3221-1',
  'MECH3023-2'],
 ['CITS1021-2',
  'CITS2021-1',
  'CITS2023-2',
  'ENSC2224-1',
  'ENSC3211-2',
  'ENSC3213-2',
  'MATH2221-1',
  'STAT1524-1'],
 ['CITS2023-2',
  'CITS3220-2',
  'CITS3223-2',
  'CITS3021-1',
  'CITS3025-1',
  'MATH3224-1',
  'MATH3221-1',
  'STAT3264-1',
  'STAT3260-2'],
 ['BIOC3221-1',
  'BIOC3224-1',
  'BIOC3225-2',
  'BIOC3225-2',
  'CITS1023-2',
  'CITS2021-1',
  'CITS2023-2',
  'MATH1211-1'],
 ['CITS2021-1',
  'CITS2023-2',
  'GRMN2025-1',
  'GRMN2020-2',
  'MATH1214-1',
  'MATH2521-2',
  'PHYS2221-1',
  'PHYS2223-2'],
 ['CITS1221-1',
  'CITS1021-1',
  'CITS1023-2',
  'CITS2023-2',
  'MGMT1135-1',
  'MGMT1136-2',
  'STAT1024-1',
  'STAT2023-2'],
 ['CITS1221-2',
  'CITS2023-2',
  'CITS3223-2',
  'CITS3021-1',
  'CITS3025-1',
  'CLAN3224-1',
  'STAT2021-1',
  'STAT3260-2'],
 ['CITS2
```

nospace_list=[]

with open(DATAFILE,'r') as file:
    for content in pattern_lists:
        newcontent=[x.replace(' ','')for x in content]
        nospace_list.append(newcontent)
    print(nospace_list)


```python

```

#### Putting it all together

Congratulations! You have now "_wrangled_" your data from its 'raw' format in the file, into a very useable form. Let's put what we have so far together in a function. 

Time for our first real marked question.

#### Checked Solution [1 point]

- To consolidate this section, turn your code so far into a function `get_patterns(filename)` that takes a filename (the input), and returns a list of lists containing all the patterns in the file, with whitespace removed (the output).

You can assume the file is the same format as the example provided, though it may have different size and different entries.

`print(get_patterns(DATAFILE)[:5])` (with DATAFILE as above) should print the same output as above.


<!-- grade_id: d2098c, points: 0, grade: false, solution: true -->
```python
def get_patterns(filename):
    pattern_lists=[]
    pattern_strings =[]
    nospace_list=[]

    with open(filename,'r') as file:
         data=file.readlines()

         for row in data:
             unit_data=row.partition('\t')
             unit_only=unit_data[0].split('+')
             pattern_lists.append(unit_only)
         for content in pattern_lists:
             newlist=[] #individual list with no space
             for units in content:
                    newcontent=units.strip()
                    newlist.append(newcontent)
             nospace_list.append(newlist)
         return nospace_list

        #newcontent=[x.replace(' ','')for x in content]
        #nospace_list.append(newcontent)
    #return nospace_list


#print(get_patterns(DATAFILE)[:5])

```

As usual, test your code on some examples. Try and 'break' your code.

Before a final test of your code, it is worth restarting the kernel from the "Kernel" drop-down menu, and re-running your code. This will clean up any leftover 'junk' in the memory from running earlier versions of bits of code.

You may find, for example, that after you fixed or changed a piece of code, you forgot to change a variable name, but the code still worked because an old value for the variable was still in memory. (This is one of the disadvantages of interpreted programming.)

* Restart the kernel and test your code again to make sure it is doing everything correctly.

Then go ahead and try the sample tests.


<!-- grade_id: get_patterns, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
DATA = "unit-patterns.txt"
assert_equal(get_patterns(DATA)[0],['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1'])
assert_equal(get_patterns(DATA)[5],['CITS2021-1', 'CITS2023-2', 'GRMN2025-1', 'GRMN2020-2', 'MATH1214-1', 'MATH2521-2', 'PHYS2221-1', 'PHYS2223-2'])
print('So far, so good! Additional tests will be applied.')

```

## Data Inspection and Interpretation

Now let's have a closer look at the data. The file contains all the study plans corresponding to each student. Let's start by working out how many *different* or *unique* plans there are.

To make it easier to see what is going on, start by defining a better print function.

* Write a function `print_patterns (patterns, numlines)` that prints the first `numlines` patterns each on a new line.

For example:

print_patterns(get_patterns(DATAFILE), 8)

```
['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1']
['CITS2023-2', 'ENSC3221-2', 'ENSC3224-1', 'ENSC3225-1', 'ENSC3227-2', 'ENSC3210-1', 'ENSC3221-1', 'MECH3023-2']
['CITS1021-2', 'CITS2021-1', 'CITS2023-2', 'ENSC2224-1', 'ENSC3211-2', 'ENSC3213-2', 'MATH2221-1', 'STAT1524-1']
['CITS2023-2', 'CITS3220-2', 'CITS3223-2', 'CITS3021-1', 'CITS3025-1', 'MATH3224-1', 'MATH3221-1', 'STAT3264-1', 'STAT3260-2']
['BIOC3221-1', 'BIOC3224-1', 'BIOC3225-2', 'BIOC3225-2', 'CITS1023-2', 'CITS2021-1', 'CITS2023-2', 'MATH1211-1']
['CITS2021-1', 'CITS2023-2', 'GRMN2025-1', 'GRMN2020-2', 'MATH1214-1', 'MATH2521-2', 'PHYS2221-1', 'PHYS2223-2']
['CITS1221-1', 'CITS1021-1', 'CITS1023-2', 'CITS2023-2', 'MGMT1135-1', 'MGMT1136-2', 'STAT1024-1', 'STAT2023-2']
['CITS1221-2', 'CITS2023-2', 'CITS3223-2', 'CITS3021-1', 'CITS3025-1', 'CLAN3224-1', 'STAT2021-1', 'STAT3260-2']
```

```python
x = [1,1,2,2,3,3]
y = set(x)
y
```

#to get unique pattern

def print_patterns(patterns,numlines):
    lst = []
    for line in patterns:
        if line not in lst:
            lst.append(line)
    print(lst[:numlines])

print_patterns(get_patterns(DATAFILE),8)

#get each line in turn 

```python
def print_patterns(patterns,numlines):
       for line in patterns[:numlines]:
             print(line)

#print_patterns(get_patterns(DATAFILE),8)



```

### Sorting the patterns

#### Checked Solution [1 point]

* Write a function `sort_patterns(patterns)` that takes a pattern list, and returns the list in which:
  * the units in each list (pattern) are sorted alpha-numerically
  * the lists (patterns) themselves are sorted alpha-numerically
    

The following test code:

```
print_patterns(get_patterns(DATA),3)
print()
print_patterns(sort_patterns(get_patterns(DATA)),3)
```

should produce the following output:

```
['CITS2224-1', 'CITS2023-2', 'CITS3221-2', 'CITS3224-1', 'CITS3225-1']
['CITS2023-2', 'ENSC3221-2', 'ENSC3224-1', 'ENSC3225-1', 'ENSC3227-2', 'ENSC3210-1', 'ENSC3221-1', 'MECH3023-2']
['CITS1021-2', 'CITS2021-1', 'CITS2023-2', 'ENSC2224-1', 'ENSC3211-2', 'ENSC3213-2', 'MATH2221-1', 'STAT1524-1']

['AACE1223-2AA', 'ACCT1121-1', 'CARS1223-2AA', 'CITS1021-1', 'CITS2023-2', 'ECON1121-1', 'FINA1221-2', 'INDG1223-2AA', 'MGMT1135-2', 'MKTG1225-2', 'STAT1524-1']
['AACE1223-2AA', 'CARS1223-2AA', 'CITS1021-1', 'CITS1023-2', 'CITS2023-2', 'INDG1223-2AA', 'LAWS1113-2', 'PSYC1121-1', 'PSYC1123-2', 'SCIE1121-1', 'STAT1024-1']
['AACE1223-2AA', 'CARS1223-2AA', 'CITS1021-2', 'CITS1023-2', 'COMM1223-2', 'INDG1223-2AA', 'STAT1023-2']```
```

_Caution: If you are using python's built-in sorting methods, it is recommended that you use `sorted` to return a new sorted list, rather than `sort` to sort the list in-place. Methods that alter a data structure 'in-place' can lead to confusing behaviour when debugging._



```python
print_patterns(get_patterns(DATA),3)
```

#wrong trial 

def sort_patterns (patterns):
    for line in patterns:
        for each_unit in line:
              first_sort=sorted(each_unit)
    return first_sort

print_patterns(get_patterns(DATA),1)
print()
print_patterns(sort_patterns(get_patterns(DATA)),10)


def sort_patterns (patterns):
    unit_sorted=[]
    #pattern_sorted=[]
    for each_list in patterns:
        sorted_list=sorted(each_list)
    print(unit_sorted)



print_patterns(sort_patterns(get_patterns(DATAFILE)),3)

<!-- grade_id: 36be71, points: 0, grade: false, solution: true -->
```python
def sort_patterns (patterns):
    unit_sorted=[]
    for each_list in patterns:
        sorted_list=sorted(each_list)
        unit_sorted.append(sorted_list)
    return sorted(unit_sorted) #sort the full list

#sort_patterns(get_patterns(DATAFILE))

#print_patterns(get_patterns(DATA),1)
#print()
#print_patterns(sort_patterns(get_patterns(DATAFILE)),1)
```

Again, test your own code, don't rely on the sample tests below.

<!-- grade_id: sort_patterns, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
DATA = "unit-patterns.txt"

sorted_patterns = sort_patterns(get_patterns(DATA))
assert_equal(sorted_patterns[0],['AACE1223-2AA', 'ACCT1121-1', 'CARS1223-2AA', 'CITS1021-1', 'CITS2023-2', 'ECON1121-1', 'FINA1221-2', 'INDG1223-2AA', 'MGMT1135-2', 'MKTG1225-2', 'STAT1524-1'])
assert_equal(sorted_patterns[5][0],'AACE1224-ACE-2')
assert_equal(sorted_patterns[10][-1],'STAT1523-2')
print("So far, so good. Additional tests will be applied.")

```

### Removing duplicates

* Write a function `remove_duplicates (patterns)` that returns a sorted list of patterns with any duplicate pattern sets removed.

_Tip: It is easier to remove duplicates if you sort the patterns first. (Why?)_

You may make use of previously defined functions.

How many *unique* study patterns are there in the class?



```python
def remove_duplicates(patterns):
    lst=[]
    for line in patterns:
        if line not in lst:
            lst.append(line)
    return lst

#remove_duplicates(sort_patterns(get_patterns(DATAFILE)))
```

- Print the percentage (to one decimal place) of the total number of study plans in the class that are different.

You should find that it is indeed a diverse group!

```python

```

## Putting it all Together

If we were going to go on and use the data, it would be useful to have a single function that retrieves and cleans the data from the file.

#### Checked Solution [1 point]

* Write a function `get_unique_patterns (filename)` that takes a filename in the above format, and returns a pair `(unique_patterns, num_patterns)` where:
  * `unique_patterns` contains a sorted list of sorted patterns,  with any duplicate patterns removed
  * `num_patterns` contains the number of unique patterns

Your code may (in fact should!) make use of functions you have already implemented.



<!-- grade_id: 729dcd, points: 0, grade: false, solution: true -->
```python
def get_unique_patterns (filename):
    x=remove_duplicates(sort_patterns(get_patterns(filename)))
    y=len(x)
    return x,y

get_unique_patterns(DATAFILE)
```

<!-- grade_id: get_unique_patterns, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
DATA = "unit-patterns.txt"
(patterns, num_patterns) = get_unique_patterns(DATA)
assert_equal(patterns[0],['AACE1223-2AA', 'ACCT1121-1', 'CARS1223-2AA', 'CITS1021-1', 'CITS2023-2', 'ECON1121-1', 'FINA1221-2', 'INDG1223-2AA', 'MGMT1135-2', 'MKTG1225-2', 'STAT1524-1'])
assert_equal(patterns[5][0],'AACE1224-ACE-2')
assert_equal(patterns[10][-1],'STAT1523-2')
print("So far, so good.")

```

## Examining the range of disciplines

For this part, it is left to you to think about the best way to break the problems down. You won't want to write new code for each different discipline, so think about how you can use structured programming to efficiently answer the questions.

* What proportion of study patterns contain: 
  * a maths unit?
  * a stats unit?
  * a business unit?
  * a medicine unit?
  * a law unit?
  * a psychology unit?
  * a music unit?
  * a philosophy unit?
  * a service learning unit?


* How many different disciplines are studied altogether by the class (for the purposes of this question assume the first four characters of the unit code represent the discipline.)

### Challenge

* Print a list of each unit that is being studied, and how many students are studying it, in order from the unit with the most students to the unit with the least.

Hints: Use a [dictionary](https://docs.python.org/3/library/stdtypes.html#mapping-types-dict), look at the `get` method, and consider the `key` argument of the inbuilt function `sorted`.

(My code is 7 lines.)

* How many units are being studied altogether?

&copy; Cara MacNish, UWA